Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```python
NAME = "Maria Baba"
STUDENT_ID = "14201089"
```

---

# Pandas and data-linkage

```python
import pandas as pd
import numpy as np
```

## Basic Pandas functionality

#### Exercise 1: Basic Series Functionality
Create a pandas `Series` with values `[1, 3, 5, np.nan, 6, 8]` and display basic functionality like `describe()`,` count()`, `sum()` etc.

<!-- grade_id: cell-f561f052eef9823b, points: 0, grade: false, solution: true -->
```python
# Exercise 1
# Your code here
# series_1 = ...

# YOUR CODE HERE
values= ([1,3, 5, np.nan, 6, 8])
series_1= pd.Series(values)
print(series_1)

d= series_1.describe()
c= series_1.count()
s= series_1.sum()
print(d, c, s)

```

Output:
```
0    1.0
1    3.0
2    5.0
3    NaN
4    6.0
5    8.0
dtype: float64
count    5.000000
mean     4.600000
std      2.701851
min      1.000000
25%      3.000000
50%      5.000000
75%      6.000000
max      8.000000
dtype: float64 5 23.0
```

<!-- grade_id: cell-33d19898476ae3ea, points: 1, grade: true, solution: false -->
```python
# Test for Exercise 1
assert series_1.sum() == 23, "Check your series values."
```

#### Exercise 2: Series with Custom Index
Create a `Series` with values `[30, 35, 40]` and indices `['Alice', 'Bob', 'Charlie']`.

<!-- grade_id: cell-5f1500ea9ce9c611, points: 0, grade: false, solution: true -->
```python
# Exercise 2
# Your code here
# series_2 = ...

# YOUR CODE HERE
values_2= ([30, 35, 40])
indices=['Alice', 'Bob', 'Charlie']
series_2= pd.Series(values_2, index=indices)
print (series_2)

```

Output:
```
Alice      30
Bob        35
Charlie    40
dtype: int64
```

<!-- grade_id: cell-9ac907bb27aa63c0, points: 1, grade: true, solution: false -->
```python
# Test for Exercise 2
assert 'Bob' in series_2, "Ensure your series has the correct index."
```

#### Exercise 3: Selection in Series
Select and print the age of Bob from the `Series` created in Exercise 2.

<!-- grade_id: cell-b320f5e1176aa070, points: 0, grade: false, solution: true -->
```python
# Exercise 3
# Your code here
age_bob = series_2.loc['Bob']
print (age_bob)
```

Output:
```
35
```

<!-- grade_id: cell-7a162b3d00a7c9dd, points: 1, grade: true, solution: false -->
```python
# Test for Exercise 3
assert age_bob == 35, "Check the value you extracted for Bob's age."
```

#### Exercise 4: Filtering Condition
Filter and display elements in `series_2` that are greater than 35.

<!-- grade_id: cell-db9ac814f6206b45, points: 0, grade: false, solution: true -->
```python
# Exercise 4
# Your code here
filtered_series = series_2.loc[lambda x: (x >35)]
print(filtered_series)
```

Output:
```
Charlie    40
dtype: int64
```

<!-- grade_id: cell-e1b0f3aabc7c001d, points: 2, grade: true, solution: false -->
```python
# Test for Exercise 4
assert filtered_series['Charlie'] == 40, "Check your filtering condition."
```

#### Exercise 5: Creating a DataFrame
Create a DataFrame using the dictionary below and assign it to a variable named `df`.

<!-- grade_id: cell-6bf2866619f54dc1, points: 0, grade: false, solution: true -->
```python
data = {
    'Age': [24, 27, 30, 35],
    'Name': ['Alice', 'Bob', 'Charlie', 'Diana']
}

# Exercise 5
# Your code here
df = pd.DataFrame(data)
display(df)

```

Output:
```
Age     Name
0   24    Alice
1   27      Bob
2   30  Charlie
3   35    Diana
```

<!-- grade_id: cell-1df22909ff854e72, points: 3, grade: true, solution: false -->
```python
# Test for Exercise 5
assert 'Name' in df, "Ensure your DataFrame contains the correct columns."
```

#### Exercise 6: Retrieving the Index or Columns
Retrieve and print the index and columns of the `DataFrame` created in Exercise 5.

<!-- grade_id: cell-024f359d866c8a29, points: 0, grade: false, solution: true -->
```python
# Exercise 6
# Your code here
index_df = df.index
print (index_df)
columns_df = df.columns
print(columns_df)
```

Output:
```
RangeIndex(start=0, stop=4, step=1)
Index(['Age', 'Name'], dtype='object')
```

<!-- grade_id: cell-00da66dca230f648, points: 1, grade: true, solution: false -->
```python
# Test for Exercise 6
assert len(index_df) == 4 and len(columns_df) == 2, "Check your index and columns extraction."
```

#### Exercise 7: Data Extraction with loc, iloc, and []
Extract and print the `Age` of `Alice` using `loc`, `iloc`, and `[]` from the `DataFrame` created in Exercise 5.

<!-- grade_id: cell-ca64155f67b2dc58, points: 0, grade: false, solution: true -->
```python
# Exercise 7
# Your code here
age_alice_loc = df.loc[0, 'Age']
print('age alice using loc:',age_alice_loc)
age_alice_iloc = df.iloc[0,0]
print('age alice using iloc:',age_alice_iloc)
age_alice_bracket= df[df['Name'] == 'Alice']['Age'].values[0]
print('age alice using brackets:',age_alice_bracket)
```

Output:
```
age alice using loc: 24
age alice using iloc: 24
age alice using brackets: 24
```

<!-- grade_id: cell-1d6147ef13c9d061, points: 2, grade: true, solution: false -->
```python
# Test for Exercise 7
assert age_alice_loc == 24 and age_alice_iloc == 24 and age_alice_bracket == 24, "Check your data extraction methods."
```

## Health Datasets

Cool! Now that we have learned how pandas works, let's start using it with some data. 

### Exercise: Deterministic Data Linkage

You will practice deterministic data linkage using synthetic datasets: `PatientDemo.csv` and `PatientVisits.csv`. Link these datasets using the `PatientID` field and perform analysis on the linked data.


```python
# Load the datasets
import pandas as pd
patient_df = pd.read_csv('PatientDemo.csv')
visits_df = pd.read_csv('PatientVisits.csv')
```

Now let's explore the data that we have to see if we can find how to link both datasets.

<!-- grade_id: cell-e5338aee544ec758, points: 0, grade: false, solution: true -->
```python
display(patient_df.head())
display(visits_df.head())

number_of_patients = patient_df['PatientID'].count()
print(f"Number of patients:{number_of_patients}")
number_of_visits = visits_df['VisitID'].count()
print(f"Number of visits:{number_of_visits}")
```

Output:
```
PatientID  FirstName LastName DateOfBirth Gender  ZipCode
0          1      Megan    Mccoy  1968-01-28      F    34536
1          2  Katherine    Bruce  1942-02-11      F    73685
2          3     Robert  Sanchez  1955-05-30      M    41751
3          4   Jonathan   Dennis  1985-07-04      F    48590
4          5    William   Wilson  1948-07-01      F    74880
   VisitID  PatientID   VisitDate DiagnosisCode TreatmentCode
0      101          1  2022-01-29           D90           T36
1      102          2  2023-05-09           D29           T18
2      103          3  2022-06-27            D6           T31
3      104          4  2023-01-14           D74           T98
4      105          5  2022-11-21           D82           T62
Number of patients:100
Number of visits:100
```

<!-- grade_id: cell-e65882382c579a81, points: 1, grade: true, solution: false -->
```python

```

### Data Linkage
Perform deterministic data linkage by merging the `patient_df` and `visits_df` DataFrames using the `PatientID` field. Store the result in a new DataFrame called `merged_df`.
```python


<!-- grade_id: cell-fac7892bca5e74e7, points: 0, grade: false, solution: true -->
```python
merged_df = pd.merge (patient_df, visits_df, on="PatientID")
display(merged_df.head())
```

Output:
```
PatientID  FirstName LastName DateOfBirth Gender  ZipCode  VisitID  \
0          1      Megan    Mccoy  1968-01-28      F    34536      101   
1          2  Katherine    Bruce  1942-02-11      F    73685      102   
2          3     Robert  Sanchez  1955-05-30      M    41751      103   
3          4   Jonathan   Dennis  1985-07-04      F    48590      104   
4          5    William   Wilson  1948-07-01      F    74880      105   

    VisitDate DiagnosisCode TreatmentCode  
0  2022-01-29           D90           T36  
1  2023-05-09           D29           T18  
2  2022-06-27            D6           T31  
3  2023-01-14           D74           T98  
4  2022-11-21           D82           T62
```

### Questions

1. How many records are in the `merged_df` DataFrame?
2. What is the average age of patients in `merged_df`?


<!-- grade_id: cell-ce9bb9e6c1a05c40, points: 0, grade: false, solution: true -->
```python
# Question 1: How many records are in the merged_df DataFrame?

num_records = merged_df['PatientID'].count()
print (f'Number of records is {num_records}')

# Question 2: What is the average age of patients in merged_df?

# Hint: You can use the datetime module to get the current year.
from datetime import datetime
current_year = datetime.now().year

merged_df['DateOfBirth']= pd.to_datetime(merged_df['DateOfBirth'])
ages= []
for x in merged_df['DateOfBirth']:
    year= int(x.year)
    age= current_year-year
    ages.append(age)
average_age= np.mean(ages)
print (f'Average age of patients is {average_age}')
```

Output:
```
Number of records is 100
Average age of patients is 52.22
```

<!-- grade_id: cell-4cb0b75dd98cd98c, points: 2, grade: true, solution: false -->
```python

```

# Probabilistic Data Linkage

Probabilistic data linkage is a technique used to bring together records from different datasets that do not share a unique identifier but have other fields in common. Unlike deterministic linkage, which requires exact matches on shared fields, probabilistic linkage considers the likelihood that two records refer to the same entity based on the similarity between these shared fields. This approach is particularly useful when working with data that may contain errors, variations, or inconsistencies in how information is recorded.

## Challenges in Data Linkage

- **Data Quality:** Datasets often have missing, misspelled, or inconsistently formatted data.
- **Variability:** The same entity might be represented slightly differently in different datasets, due to typos, abbreviations, or variations in how data is entered.
- **Absence of Unique Identifiers:** There might not be a unique, consistent identifier shared across datasets.

## Approach

Probabilistic linkage typically involves the following steps:

1. **Selecting Variables:** Choosing which fields (or combinations of fields) will be used to link records between datasets.
2. **Measuring Similarity:** Calculating a similarity score between records based on the selected variables.
3. **Setting Thresholds:** Determining a similarity threshold above which records will be considered a match.

Various similarity or distance metrics can be used in step 2, depending on the nature of the data and the specific requirements of the linkage task. One such metric, designed to handle some of the challenges mentioned above, is the Jaro-Winkler Distance, which you have seen in the lecture about data linkage.


## Understanding Jaro-Winkler Distance

The Jaro-Winkler Distance (JWD) is a measure of similarity between two strings. It is a variant of the Jaro distance metric and is mainly used in the area of record linkage. The Jaro-Winkler Distance metric is designed to capture similarity between two strings while accounting for possible errors such as typos and characters out of place.

**Mathematical Definition**

The Jaro-Winkler Distance between two strings $ s1 $ and $ s2 $ is calculated as follows:

1. **Jaro Distance Calculation:**
    - Let $ m $ be the number of matching characters between the two strings.
    - Let $ t $ be the number of transpositions between the two strings.
    - The Jaro distance $ d_j $ is then given by the formula:
    $$
    d_j = \frac{1}{3} \left( \frac{m}{|s1|} + \frac{m}{|s2|} + \frac{m-t}{m} \right)
    $$
    where $ |s1| $ and $ |s2| $ are the lengths of the strings $ s1 $ and $ s2 $ respectively.

2. **Jaro-Winkler Distance Calculation:**
    - Let $ l $ be the length of common prefix at the start of the string (maximum 4 characters).
    - The Jaro-Winkler distance $ d_{jw} $ is then given by the formula:
    $$
    d_{jw} = d_j + l p (1 - d_j)
    $$
    where $ p $ is a constant scaling factor (usually set to $ 0.1 $).

**Interpretation**

- The Jaro-Winkler Distance ranges from $ 0 $ to $ 1 $, where $ 0 $ represents completely dissimilar strings and $ 1 $ represents identical strings.
- The distance is symmetric, meaning $ d_{jw}(s1, s2) = d_{jw}(s2, s1) $.
- It is particularly useful for short strings and for applications where the strings being compared have small length variations, making it widely used in record linkage tasks.
- The Jaro-Winkler adjustment gives more favorable ratings to strings that match from the beginning, making it useful for cases where the position of characters in the string is important.

**Example**

For instance, the strings "MARTHA" and "MARHTA" have a Jaro distance of approximately $ 0.944 $ and a Jaro-Winkler Distance of approximately $ 0.961 $ with a prefix length of $ 3 $ and scaling factor $ p = 0.1 $.


## Exercise: Implement distance

Complete the function below by implementing the Jaro-Winkler Distance Calculation.


<!-- grade_id: cell-32fed749c259b2fb, points: 0, grade: false, solution: true -->
```python
def jaro_winkler_distance(s1: str, s2: str) -> float:
    if not s1 or not s2:
        return 0.0

    m = 0  # Number of matching characters
    t = 0  # Number of transpositions
    l = min(len(s1), len(s2))  # Length of the shorter string

    # Compute the number of matching characters and transpositions
    for i in range(l):
        if s1[i] == s2[i]:
            m += 1
        elif s1[i] != s2[i] and (i == 0 or (i > 0 and s1[i - 1] == s2[i - 1])):
            t += 1

    # Compute the Jaro distance
    jaro = (1/3) * ((m / len(s1)) + (m / len(s2)) + ((m - t / 2) / m)) if m != 0 else 0.0
    
    # Compute the Jaro-Winkler distance
    p = 0.1  # Scaling factor (constant)
    l = 0  # Length of common prefix at the start of the string (max 4)
    for i in range(min(len(s1), len(s2), 4)):
        if s1[i] == s2[i]:
            l += 1
        else:
            break
    result= jaro + l*p*(1-jaro)
    return result
```

<!-- grade_id: cell-3583a3edc1c47026, points: 2, grade: true, solution: false -->
```python
assert jaro_winkler_distance('12', '1') == 0.85, 'Check your implementation.'
assert jaro_winkler_distance("SAME", "SAME") == 1.0, \
    f'Expected 1.0 for identical strings, but got {jaro_winkler_distance("SAME", "SAME")}'
assert jaro_winkler_distance("", "NOTSAME") == 0.0, \
    f'Expected 0.0 for empty string comparison, but got {jaro_winkler_distance("", "NOTSAME")}'
```

# Dataset Introduction

In this exercise, we will work with two synthetic datasets: `EmploymentData.csv` and `SocialNetworkData.csv`. These datasets represent a common scenario in data linkage projects where we have information about individuals spread across different sources.

## EmploymentData.csv

The `EmploymentData.csv` dataset contains information about employees in various companies. Each record provides details about an employee's first name (`EmployeeFirstName`), last name (`EmployeeLastName`), the company they work for (`Company`), their position (`Position`), and their employment start date (`StartDate`).

Here's a preview of the `EmploymentData.csv` structure:

| EmployeeFirstName | EmployeeLastName | Company  | Position                | StartDate  |
|-------------------|------------------|----------|-------------------------|------------|
| John              | Doe              | ABC Corp | Data Scientist          | 2022-01-15 |
| ...               | ...              | ...      | ...                     | ...        |

## SocialNetworkData.csv

The `SocialNetworkData.csv` dataset represents profiles from a professional social network. Each record includes the first name (`FirstName`), last name (`LastName`), current job title (`CurrentJobTitle`), the number of connections (`ConnectionsCount`), and the profile creation date (`ProfileCreationDate`).

Here's a preview of the `SocialNetworkData.csv` structure:

| FirstName | LastName | CurrentJobTitle       | ConnectionsCount | ProfileCreationDate |
|-----------|----------|-----------------------|------------------|---------------------|
| Jon       | Does     | Senior Data Scientist | 300              | 2018-06-05          |
| ...       | ...      | ...                   | ...              | ...                 |

## Objective

Your task is to link records between these datasets probabilistically, based on the similarity of names using the Jaro-Winkler distance measure you will implement. Due to potential variations in how names are represented in each dataset (e.g., nicknames, typos), the linkage process requires careful consideration and application of string similarity measures.

Load the `EmploymentData.csv` and `SocialNetworkData.csv` datasets into two separate DataFrames: `employment_df` and `network_df`.

```python
# Load the employment data
employment_df = pd.read_csv('EmploymentData.csv')

# Load the social network data
network_df = pd.read_csv('SocialNetworkData.csv')

```

### Probabilistic Data Linkage
Use the `jaro_winkler_distance` function you implemented to link records between the `employment_df` and `network_df` DataFrames with the highest average similarity and a threshold of above 0.40. Store the result in a new DataFrame called `linked_df`.



<!-- grade_id: cell-77d2d30109393775, points: 0, grade: false, solution: true -->
```python
threshold = 0.65
linked_records = []

for idx1, row1 in employment_df.iterrows():
    highest_similarity = 0
    best_match_index = -1
    for idx2, row2 in network_df.iterrows():
        first_name_similarity = jaro_winkler_distance(row1['EmployeeFirstName'], row2['FirstName'])
        last_name_similarity = jaro_winkler_distance(row1['EmployeeLastName'], row2['LastName'])
        average_similarity = (first_name_similarity + last_name_similarity) / 2
        if average_similarity > highest_similarity:
            data_row = row2
            highest_similarity = average_similarity
            best_match_index = idx2
    if highest_similarity >= threshold:
        linked_records.append(highest_similarity)
        # Save the record pair if the similarity is above the threshold
        # and add the similarity value to the record pair

linked_df= pd.DataFrame(linked_records)
```

<!-- grade_id: cell-d905cbbfeaa23aad, points: 0, grade: true, solution: false -->
```python
linked_df.head()
```

Output:
```
0
0  1.000000
1  1.000000
2  0.941667
3  0.718750
4  1.000000
```

## Evaluation of Data Linkage Results

After performing data linkage, it's crucial to evaluate the quality of our matches. This step helps identify potential errors and assess the effectiveness of our linkage technique.

To quantify the quality, we'll focus on two key metrics:
- **Precision**: Measures the correctness of the matches. A higher precision means fewer false positives.
- **Recall**: Measures the completeness of the matches. A higher recall means fewer true matches were missed.

Precision and recall are defined as:
$$ \text{Precision} = \frac{\text{True Positives (TP)}}{\text{True Positives (TP) + False Positives (FP)}} $$
$$ \text{Recall} = \frac{\text{True Positives (TP)}}{\text{True Positives (TP) + False Negatives (FN)}} $$

Using the ground truth (matches based on indices in this synthetic example), we can calculate these metrics to understand how well our linkage algorithm performed.


```python
def evaluate_matches(employment_df, network_df):
    # Ground truth matches based on indices
    ground_truth = set([(idx, idx) for idx in employment_df.index if idx in network_df.index])

    # Matches made by the linkage algorithm
    predicted_matches = set([(idx1, idx2) for idx1, row1 in employment_df.iterrows() 
                            for idx2, row2 in network_df.iterrows() 
                            if (row1['EmployeeFirstName'], row1['EmployeeLastName']) == 
                                (row2['FirstName'], row2['LastName'])])

    # True Positives: Ground truth matches that are in the predicted matches
    TP = len(ground_truth.intersection(predicted_matches))

    # False Positives: Predicted matches that are not in the ground truth
    FP = len(predicted_matches.difference(ground_truth))

    # False Negatives: Ground truth matches that are not in the predicted matches
    FN = len(ground_truth.difference(predicted_matches))

    # Calculate Precision and Recall
    precision = TP / (TP + FP) if TP + FP > 0 else 0
    recall = TP / (TP + FN) if TP + FN > 0 else 0

    return precision, recall

result= evaluate_matches(employment_df, network_df)
print(result)

```

Output:
```
(1.0, 0.73)
```

### Set threshold
Now play around with the threshold and see how you could optimize the matching.

<!-- grade_id: cell-9098da9f7c3003d3, points: 0, grade: false, solution: true -->
```python
def evaluate_matches(employment_df, network_df):
    # Ground truth matches based on indices
    ground_truth = set([(idx, idx) for idx in employment_df.index if idx in network_df.index])

    # Matches made by the linkage algorithm
    predicted_matches = set()
    threshold = 0.30

    for idx1, row1 in employment_df.iterrows():
        highest_similarity = 0
        best_match_index = -1
        for idx2, row2 in network_df.iterrows():
            first_name_similarity = jaro_winkler_distance(row1['EmployeeFirstName'], row2['FirstName'])
            last_name_similarity = jaro_winkler_distance(row1['EmployeeLastName'], row2['LastName'])
            average_similarity = (first_name_similarity + last_name_similarity) / 2
            if average_similarity > highest_similarity:
                data_row = row2
                highest_similarity = average_similarity
                best_match_index = idx2
        if highest_similarity >= threshold:
            predicted_matches.add((idx1, best_match_index))
    # True Positives: Ground truth matches that are in the predicted matches
    TP = len(ground_truth.intersection(predicted_matches))

    # False Positives: Predicted matches that are not in the ground truth
    FP = len(predicted_matches.difference(ground_truth))

    # False Negatives: Ground truth matches that are not in the predicted matches
    FN = len(ground_truth.difference(predicted_matches))

    # Calculate Precision and Recall
    precision = TP / (TP + FP) if TP + FP > 0 else 0
    recall = TP / (TP + FN) if TP + FN > 0 else 0

    return precision, recall

result= evaluate_matches(employment_df, network_df)
print(result)

```

Output:
```
(0.98, 0.98)
```

## Conceptual Exercise: Data Linkage with Different Measurement Scales

### Background
Imagine you are working with two datasets collected from two different social surveys conducted in a city:

1. **CityResidentSurvey.csv:**
   - **Age:** Age of the resident.
   - **Income:** Monthly income of the resident in USD.
   - **EducationLevel:** Highest level of education attained (coded as 1: High School, 2: Bachelor’s, 3: Master’s, 4: Doctorate).
   - **ResidentID:** A unique identifier for each resident in the survey.

2. **NeighborhoodWellbeingSurvey.csv:**
   - **Resident_Age:** Age of the resident, but with a ±2 years error margin due to the way it was collected.
   - **Annual_Income:** Yearly income of the resident in USD.
   - **Edu_Level:** Highest level of education attained, described with words (High School, Bachelor's Degree, Master's Degree, Doctorate).
   - **WellbeingIndex:** A score representing the resident’s perceived wellbeing.
   - **Resident_ID:** A unique identifier, but it does not match the ResidentID in the CityResidentSurvey.

### Task
You are tasked to link records between these two datasets. However, the data points are not measured exactly the same way in both datasets. Specifically:
- Age is measured with a ±2 years error in the second dataset.
- Income is reported monthly in the first dataset and annually in the second.
- Education levels are coded numerically in the first dataset and described with words in the second.

### Question
Describe a step-by-step approach to link records between the two datasets as accurately as possible. Consider the differences in the measurement scales and potential errors in the data. How would you account for these differences to improve the accuracy of your linkage? You don't have to write any code, just explain your approach conceptually.


<!-- grade_id: cell-8ba6fc8b8431eecc, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE: First I would try to "fix" the way the data points are measured so they can be represented in the same way in both datasets. Looking at the Neighborhood Well Being survey, I would divide the income by 12 (number of months) to get a monthly income as in the other survey. Similarly, I would convert the level of education from strings to integers (or codes) as shown in the first survey, to make measurments similar. In terms of age, I would try to find an average between the ages showed in the first and second dataset, to get an approximate understanding of what the age can be.

Then, I would link the datasets based on education as it would be the most accurate measurment.
